Create a variable named `a` with value `5`, a variable `b` with value `"hello"`, and a variable `c` with value `[1, 2, "test"]`.

<!-- grade_id: soln, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_simple, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(a)) == "<class 'int'>", "type of a is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert str(a) == "5", "value of a is not correct"

assert str(type(b)) == "<class 'str'>", "type of b is not str. b should be an str"
assert str(len(b)) == "5", "length of b is not correct"
assert str(b.lower()) == "hello", "value of b is not correct"
assert str(b) == "hello", "correct string value of b but incorrect case of letters"

assert str(type(c)) == "<class 'list'>", "type of c is not list. c should be a list"
assert str(len(c)) == "3", "length of c is not correct"
assert str(sorted(map(str, c))) == "['1', '2', 'test']", "values of c are not correct"
assert str(c) == "[1, 2, 'test']", "order of elements of c is not correct"

print('Success!')
```