Create a variable named `a` with value `5`, a variable `b` with value `"hello"`, and a variable `c` with value `[1, 2, "test"]`.

<!-- grade_id: soln, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_hidden, points: 1, grade: true, solution: false -->
```python

assert str(type(a)) == "<class 'int'>", "type of type(a) is not correct"

assert str(type(b)) == "<class 'str'>", "type of type(b) is not correct"

assert str(type(c)) == "<class 'list'>", "type of type(c) is not correct"

print('Success!')
```