<!-- grade_id: soln, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
a = 5
b = "hello"
c = [1, 2, "test"]
d = {1 : 6, 3: 5} #right type, wrong value
e = [5, -3.3, 'd'] #right values, wrong type
f = True # wrong value
def fun(x):
    if x < 0:
        raise ValueError
    else:
        return x
```

<!-- grade_id: test_multi, points: 1, grade: true, solution: false -->
```python
# the basic tests from the simple notebook
from hashlib import sha1
assert str(type(a)) == "<class 'int'>", "type of a is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert str(a) == "5", "value of a is not correct"

assert str(type(b)) == "<class 'str'>", "type of b is not str. b should be an str"
assert str(len(b)) == "5", "length of b is not correct"
assert str(b.lower()) == "hello", "value of b is not correct"
assert str(b) == "hello", "correct string value of b but incorrect case of letters"

assert str(type(c)) == "<class 'list'>", "type of c is not list. c should be a list"
assert str(len(c)) == "3", "length of c is not correct"
assert str(sorted(map(str, c))) == "['1', '2', 'test']", "values of c are not correct"
assert str(c) == "[1, 2, 'test']", "order of elements of c is not correct"

print('Success!')
```

<!-- grade_id: cell-f2803ba7c42d03ab, points: 1, grade: true, solution: false -->
```python
# multiple expressions per line
from hashlib import sha1
assert str(type(a)) == "<class 'int'>", "type of type(a) is not correct"

assert str(type(str(a))) == "<class 'str'>", "type of str(a) is not str. str(a) should be an str"
assert str(len(str(a))) == "1", "length of str(a) is not correct"
assert str(str(a).lower()) == "5", "value of str(a) is not correct"
assert str(str(a)) == "5", "correct string value of str(a) but incorrect case of letters"

assert str(type(a == "5")) == "<class 'bool'>", "type of a == \"5\" is not bool. a == \"5\" should be a bool"
assert str(a == "5") == "False", "boolean value of a == \"5\" is not correct"

assert str(type([ch for ch in b])) == "<class 'list'>", "type of [ch for ch in b] is not list. [ch for ch in b] should be a list"
assert str(len([ch for ch in b])) == "5", "length of [ch for ch in b] is not correct"
assert str(sorted(map(str, [ch for ch in b]))) == "['e', 'h', 'l', 'l', 'o']", "values of [ch for ch in b] are not correct"
assert str([ch for ch in b]) == "['h', 'e', 'l', 'l', 'o']", "order of elements of [ch for ch in b] is not correct"

assert str(type(len(c))) == "<class 'int'>", "type of len(c) is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert str(len(c)) == "3", "value of len(c) is not correct"

assert str(type(a)) == "<class 'int'>", "type of a is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert str(a) == "5", "value of a is not correct"


# intervening regular code
print("hello!")

# differing spacings, numbers of comment characters, trailing whitespace
assert str(type(a)) == "<class 'int'>", "type of a is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert str(a) == "5", "value of a is not correct"

assert str(type(a)) == "<class 'int'>", "type of a is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert str(a) == "5", "value of a is not correct"

assert str(type(a)) == "<class 'int'>", "type of a is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert str(a) == "5", "value of a is not correct"

assert str(type(a)) == "<class 'int'>", "type of a is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert str(a) == "5", "value of a is not correct"

assert str(type(a)) == "<class 'int'>", "type of a is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert str(a) == "5", "value of a is not correct"

assert str(type(a)) == "<class 'int'>", "type of a is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert str(a) == "5", "value of a is not correct"



print('Success!')
```

<!-- grade_id: cell-693350420ec62f1b, points: 1, grade: true, solution: false -->
```python
# a few common types
from hashlib import sha1
assert str(type(a)) == "<class 'int'>", "type of a is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert str(a) == "5", "value of a is not correct"

assert str(type(b)) == "<class 'str'>", "type of b is not str. b should be an str"
assert str(len(b)) == "5", "length of b is not correct"
assert str(b.lower()) == "hello", "value of b is not correct"
assert str(b) == "hello", "correct string value of b but incorrect case of letters"

assert str(type(c)) == "<class 'list'>", "type of c is not list. c should be a list"
assert str(len(c)) == "3", "length of c is not correct"
assert str(sorted(map(str, c))) == "['1', '2', 'test']", "values of c are not correct"
assert str(c) == "[1, 2, 'test']", "order of elements of c is not correct"

assert str(type(d)) == "<class 'dict'>", "type of d is not dict. d should be a dict"
assert str(len(list(d.keys()))) == "2", "number of keys of d is not correct"
assert str(sorted(map(str, d.keys()))) == "['3', 'a']", "keys of d are not correct"
assert str(sorted(map(str, d.values()))) == "['[1.2, 3]', 'f']", "correct keys, but values of d are not correct"
assert str(d) == "{'a': 'f', 3: [1.2, 3]}", "correct keys and values, but incorrect correspondence in keys and values of d"

assert str(type(e)) == "<class 'tuple'>", "type of e is not tuple. e should be a tuple"
assert str(len(e)) == "3", "length of e is not correct"
assert str(sorted(map(str, e))) == "['-3.3', '5', 'd']", "values of e are not correct"
assert str(e) == "(5, -3.3, 'd')", "order of elements of e is not correct"

assert str(type(f)) == "<class 'bool'>", "type of f is not bool. f should be a bool"
assert str(f) == "False", "boolean value of f is not correct"

print('Success!')
```

<!-- grade_id: cell-13479eb0e5fff152, points: 1, grade: true, solution: false -->
```python

# a function that checks whether a function throws an error
"""Check that squares raises an error for invalid inputs"""
def test_func_throws(func, ErrorType):
    try:
        func()
    except ErrorType:
        return True
    else:
        print('Did not raise right type of error!')
        return False

# test a custom function
from hashlib import sha1
assert str(type(fun(3))) == "<class 'int'>", "type of fun(3) is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert str(fun(3)) == "3", "value of fun(3) is not correct"

assert str(type(test_func_throws(lambda : fun(-4), ValueError))) == "<class 'bool'>", "type of test_func_throws(lambda : fun(-4), ValueError) is not bool. test_func_throws(lambda : fun(-4), ValueError) should be a bool"
assert str(test_func_throws(lambda : fun(-4), ValueError)) == "True", "boolean value of test_func_throws(lambda : fun(-4), ValueError) is not correct"

assert str(type(test_func_throws(lambda : fun(0), ValueError))) == "<class 'bool'>", "type of test_func_throws(lambda : fun(0), ValueError) is not bool. test_func_throws(lambda : fun(0), ValueError) should be a bool"
assert str(test_func_throws(lambda : fun(0), ValueError)) == "False", "boolean value of test_func_throws(lambda : fun(0), ValueError) is not correct"


print('Success!')
```