A small test notebook to test the removal of of hidden test cases. 

<!-- grade_id: hidetest_input, points: 0, grade: false, solution: true -->
```python
# Set a to 1 
a = 1
```

<!-- grade_id: hidetest_middle, points: 1, grade: true, solution: false -->
```python
assert a == 1
### HIDESTART 
assert a == 2
assert a == 4
assert a == 6
### HIDEEND
assert a == 3
```

<!-- grade_id: hidetest_start, points: 1, grade: true, solution: false -->
```python
### HIDESTART 
assert a == 2
assert a == 4
assert a == 6
### HIDEEND
assert a == 3
```

<!-- grade_id: hidetest_end, points: 1, grade: true, solution: false -->
```python
assert a == 1
### HIDESTART 
assert a == 2
assert a == 4
assert a == 6
### HIDEEND
```

<!-- grade_id: hidetest_only, points: 1, grade: true, solution: false -->
```python
### HIDESTART 
assert a == 2
assert a == 4
assert a == 6
### HIDEEND
```