<!-- grade_id: set_a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: foo, points: 1, grade: true, solution: false -->
```python
print("Success!")
```

<!-- grade_id: bar, points: 1, grade: true, solution: false -->
```python
#assert a == 1
```

<!-- grade_id: baz, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: quux, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```