Validating this file should fail even though the assignment has 0 points.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
def squares(n):
   # YOUR CODE HERE
    return [0]
```

<!-- grade_id: correct_squares, points: 0, grade: true, solution: false -->
```python
assert squares(1) == [1]
```