Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

```python
# import plotting libraries
%matplotlib inline
import matplotlib.pyplot as plt
```

---
## Part A

Write a function that returns a list of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by raising a `ValueError`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$. Check that it does:

```python
squares(10)
```

<!-- grade_id: correct_squares, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
from nose.tools import assert_equal
assert_equal(squares(1), [1])
assert_equal(squares(2), [1, 4])
assert_equal(squares(10), [1, 4, 9, 16, 25, 36, 49, 64, 81, 100])
assert_equal(squares(11), [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121])
```

<!-- grade_id: squares_invalid_input, points: 1, grade: true, solution: false -->
```python
"""Check that squares raises an error for invalid inputs"""
from nose.tools import assert_raises
assert_raises(ValueError, squares, 0)
assert_raises(ValueError, squares, -4)
```

---

## Part B

Using your `squares` function, write a function that computes the sum of the squares of the numbers from 1 to $n$. Your function should call the `squares` function -- it should NOT reimplement its functionality.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def sum_of_squares(n):
    """Compute the sum of the squares of numbers from 1 to n."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

The sum of squares from 1 to 10 should be 385. Verify that this is the answer you get:

```python
sum_of_squares(10)
```

<!-- grade_id: correct_sum_of_squares, points: 1, grade: true, solution: false -->
```python
"""Check that sum_of_squares returns the correct answer for various inputs."""
assert_equal(sum_of_squares(1), 1)
assert_equal(sum_of_squares(2), 5)
assert_equal(sum_of_squares(10), 385)
assert_equal(sum_of_squares(11), 506)
```

<!-- grade_id: sum_of_squares_uses_squares, points: 1, grade: true, solution: false -->
```python
"""Check that sum_of_squares relies on squares."""
orig_squares = squares
del squares
try:
    assert_raises(NameError, sum_of_squares, 1)
except AssertionError:
    raise AssertionError("sum_of_squares does not use squares")
finally:
    squares = orig_squares
```

---
## Part C

Using LaTeX math notation, write out the equation that is implemented by your `sum_of_squares` function.

<!-- grade_id: sum_of_squares_equation, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

---
## Part D

Create a plot of the sum of squares for $n=1$ to $n=15$. Make sure to appropriately label the $x$-axis and $y$-axis, and to give the plot a title. Set the $x$-axis limits to be 1 (minimum) and 15 (maximum).

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
fig, ax = plt.subplots() # do not delete this line!
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: plot_axis_limits, points: 0.1, grade: true, solution: false -->
```python
"""Check that the axis limits are correct."""
assert_equal(ax.get_xlim(), (1.0, 15.0))
```

<!-- grade_id: plot_labels, points: 0.2, grade: true, solution: false -->
```python
"""Check that the axis labels and title are set."""
assert ax.get_xlabel() != "", "xlabel not set"
assert ax.get_ylabel() != "", "ylabel not set"
assert ax.get_title() != "", "title not set"
```

<!-- grade_id: plot_data, points: 0.2, grade: true, solution: false -->
```python
"""Check that the correct data was used."""
lines = ax.get_lines()
assert_equal(len(lines), 1)

# check the x-axis data
xdata = lines[0].get_xdata()
xdata_correct = np.arange(1, 16)
assert_equal(xdata.shape, xdata_correct.shape)
assert (xdata == xdata_correct).all()

# check the y-axis data
ydata = lines[0].get_ydata()
ydata_correct = np.array([sum_of_squares(x) for x in xdata])
assert_equal(ydata.shape, ydata_correct.shape)
assert (ydata == ydata_correct).all()
```