For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
## Part A (2 points)

Write a function that returns a list of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by raising a `ValueError`.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$. Check that it does:

```python
squares(10)
```

<!-- grade_id: correct_squares, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
from hashlib import sha1
assert str(type(squares(1))) == "<class 'list'>", "type of squares(1) is not list. squares(1) should be a list"
assert str(len(squares(1))) == "1", "length of squares(1) is not correct"
assert str(sorted(map(str, squares(1)))) == "['1']", "values of squares(1) are not correct"
assert str(squares(1)) == "[1]", "order of elements of squares(1) is not correct"

assert str(type(squares(2))) == "<class 'list'>", "type of squares(2) is not list. squares(2) should be a list"
assert str(len(squares(2))) == "2", "length of squares(2) is not correct"
assert str(sorted(map(str, squares(2)))) == "['1', '4']", "values of squares(2) are not correct"
assert str(squares(2)) == "[1, 4]", "order of elements of squares(2) is not correct"

assert sha1(str(type(squares(3))).encode("utf-8")+b"72d3079359b65a39").hexdigest() == "f310d99433757a135d7fccbaf323351751681395", "type of squares(3) is not list. squares(3) should be a list"
assert sha1(str(len(squares(3))).encode("utf-8")+b"72d3079359b65a39").hexdigest() == "97a57d6cc581362a54150ea0ebdc4c35e9303bad", "length of squares(3) is not correct"
assert sha1(str(sorted(map(str, squares(3)))).encode("utf-8")+b"72d3079359b65a39").hexdigest() == "a24325644c55e1932ce636d5e537e757e3f7673e", "values of squares(3) are not correct"
assert sha1(str(squares(3)).encode("utf-8")+b"72d3079359b65a39").hexdigest() == "8706d62afb2d2ec1d32fc03581040cd795b27525", "order of elements of squares(3) is not correct"

print('Success!')
```

<!-- grade_id: squares_invalid_input, points: 1, grade: true, solution: false -->
```python
"""Check that squares raises an error for invalid inputs"""
def test_func_throws(func, ErrorType):
    try:
        func()
    except ErrorType:
        return True
    else:
        print('Did not raise right type of error!')
        return False
    
from hashlib import sha1
assert str(type(test_func_throws(lambda : squares(0), ValueError))) == "<class 'bool'>", "type of test_func_throws(lambda : squares(0), ValueError) is not bool. test_func_throws(lambda : squares(0), ValueError) should be a bool"
assert str(test_func_throws(lambda : squares(0), ValueError)) == "True", "boolean value of test_func_throws(lambda : squares(0), ValueError) is not correct"

assert str(type(test_func_throws(lambda : squares(-4), ValueError))) == "<class 'bool'>", "type of test_func_throws(lambda : squares(-4), ValueError) is not bool. test_func_throws(lambda : squares(-4), ValueError) should be a bool"
assert str(test_func_throws(lambda : squares(-4), ValueError)) == "True", "boolean value of test_func_throws(lambda : squares(-4), ValueError) is not correct"


print('Success!')
```

---

## Part B (1 point)

Using your `squares` function, write a function that computes the sum of the squares of the numbers from 1 to $n$. Your function should call the `squares` function -- it should NOT reimplement its functionality.

<!-- grade_id: sum_of_squares, points: 0, grade: false, solution: true -->
```python
def sum_of_squares(n):
    """Compute the sum of the squares of numbers from 1 to n."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

The sum of squares from 1 to 10 should be 385. Verify that this is the answer you get:

```python
sum_of_squares(10)
```

<!-- grade_id: correct_sum_of_squares, points: 0.5, grade: true, solution: false -->
```python
"""Check that sum_of_squares returns the correct answer for various inputs."""
from hashlib import sha1
assert str(type(sum_of_squares(1))) == "<class 'int'>", "type of sum_of_squares(1) is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert str(sum_of_squares(1)) == "1", "value of sum_of_squares(1) is not correct"

assert str(type(sum_of_squares(2))) == "<class 'int'>", "type of sum_of_squares(2) is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert str(sum_of_squares(2)) == "5", "value of sum_of_squares(2) is not correct"

assert str(type(sum_of_squares(10))) == "<class 'int'>", "type of sum_of_squares(10) is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert str(sum_of_squares(10)) == "385", "value of sum_of_squares(10) is not correct"

assert str(type(sum_of_squares(11))) == "<class 'int'>", "type of sum_of_squares(11) is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert str(sum_of_squares(11)) == "506", "value of sum_of_squares(11) is not correct"


print('Success!')
```

<!-- grade_id: sum_of_squares_uses_squares, points: 0.5, grade: true, solution: false -->
```python
"""Check that sum_of_squares relies on squares."""

orig_squares = squares
del squares

from hashlib import sha1
assert str(type(test_func_throws(lambda : sum_of_squares(1), NameError))) == "<class 'bool'>", "type of test_func_throws(lambda : sum_of_squares(1), NameError) is not bool. test_func_throws(lambda : sum_of_squares(1), NameError) should be a bool"
assert str(test_func_throws(lambda : sum_of_squares(1), NameError)) == "True", "boolean value of test_func_throws(lambda : sum_of_squares(1), NameError) is not correct"


squares = orig_squares

print('Success!')
```

---
## Part C (1 point)

Using LaTeX math notation, write out the equation that is implemented by your `sum_of_squares` function.

<!-- grade_id: sum_of_squares_equation, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

---
## Part D (2 points)

Find a usecase for your `sum_of_squares` function and implement that usecase in the cell below.

<!-- grade_id: sum_of_squares_application, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

---
## Part E (4 points)

State the formulae for an arithmetic and geometric sum and verify them numerically for an example of your choice.