<!-- grade_id: cell-39cee8b588a2b231, points: 0, grade: false, solution: true -->
```python
def twice(x):
    # YOUR CODE HERE
    return x+x
    raise NotImplementedError()
```

<!-- grade_id: cell-285e58af6409632f, points: 1, grade: true, solution: false -->
```python
assert twice(9)==18
```

Define a function to calculate three times a number!

```python
def thrice(x):
    return 3*x
```

```python

```