For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
## Part A (2 points)

Write a function that returns a list of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by raising a `ValueError`.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    ### BEGIN SOLUTION
    if n < 1:
        raise ValueError("n must be greater than or equal to 1")
    return [i ** 2 for i in range(1, n + 1)]
    ### END SOLUTION
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$. Check that it does:

```python
squares(10)
```

<!-- grade_id: correct_squares, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
from hashlib import sha1
assert str(type(squares(1))) == """<class 'list'>""", """type of squares(1) is not list. squares(1) should be a list"""
assert str(len(squares(1))) == """1""", """length of squares(1) is not correct"""
assert str(sorted(map(str, squares(1)))) == """['1']""", """values of squares(1) are not correct"""
assert str(squares(1)) == """[1]""", """order of elements of squares(1) is not correct"""

assert str(type(squares(2))) == """<class 'list'>""", """type of squares(2) is not list. squares(2) should be a list"""
assert str(len(squares(2))) == """2""", """length of squares(2) is not correct"""
assert str(sorted(map(str, squares(2)))) == """['1', '4']""", """values of squares(2) are not correct"""
assert str(squares(2)) == """[1, 4]""", """order of elements of squares(2) is not correct"""

assert sha1(str(type(squares(3))).encode("utf-8")+b"88ffb547eac14d93").hexdigest() == """952350db81d1d714094250ac6fe0ad5e7114490e""", """type of squares(3) is not list. squares(3) should be a list"""
assert sha1(str(len(squares(3))).encode("utf-8")+b"88ffb547eac14d93").hexdigest() == """0028b3b04351fb2909f795ae5fcba15846f9977c""", """length of squares(3) is not correct"""
assert sha1(str(sorted(map(str, squares(3)))).encode("utf-8")+b"88ffb547eac14d93").hexdigest() == """aabf7ee82847b93c3dc5903ec8235c10e03c9241""", """values of squares(3) are not correct"""
assert sha1(str(squares(3)).encode("utf-8")+b"88ffb547eac14d93").hexdigest() == """4ba1f7d938fe81de3643f77ea0fadb6b4b870c2a""", """order of elements of squares(3) is not correct"""

print('Success!')
```

<!-- grade_id: squares_invalid_input, points: 1, grade: true, solution: false -->
```python
"""Check that squares raises an error for invalid inputs"""
def test_func_throws(func, ErrorType):
    try:
        func()
    except ErrorType:
        return True
    else:
        print('Did not raise right type of error!')
        return False
    
from hashlib import sha1
assert str(type(test_func_throws(lambda : squares(0), ValueError))) == """<class 'bool'>""", """type of test_func_throws(lambda : squares(0), ValueError) is not bool. test_func_throws(lambda : squares(0), ValueError) should be a bool"""
assert str(test_func_throws(lambda : squares(0), ValueError)) == """True""", """boolean value of test_func_throws(lambda : squares(0), ValueError) is not correct"""

assert str(type(test_func_throws(lambda : squares(-4), ValueError))) == """<class 'bool'>""", """type of test_func_throws(lambda : squares(-4), ValueError) is not bool. test_func_throws(lambda : squares(-4), ValueError) should be a bool"""
assert str(test_func_throws(lambda : squares(-4), ValueError)) == """True""", """boolean value of test_func_throws(lambda : squares(-4), ValueError) is not correct"""


print('Success!')
```

---

## Part B (1 point)

Using your `squares` function, write a function that computes the sum of the squares of the numbers from 1 to $n$. Your function should call the `squares` function -- it should NOT reimplement its functionality.

<!-- grade_id: sum_of_squares, points: 0, grade: false, solution: true -->
```python
def sum_of_squares(n):
    """Compute the sum of the squares of numbers from 1 to n."""
    ### BEGIN SOLUTION
    return sum(squares(n))
    ### END SOLUTION
```

The sum of squares from 1 to 10 should be 385. Verify that this is the answer you get:

```python
sum_of_squares(10)
```

<!-- grade_id: correct_sum_of_squares, points: 0.5, grade: true, solution: false -->
```python
"""Check that sum_of_squares returns the correct answer for various inputs."""
from hashlib import sha1
assert str(type(sum_of_squares(1))) == """<class 'int'>""", """type of sum_of_squares(1) is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"""
assert str(sum_of_squares(1)) == """1""", """value of sum_of_squares(1) is not correct"""

assert str(type(sum_of_squares(2))) == """<class 'int'>""", """type of sum_of_squares(2) is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"""
assert str(sum_of_squares(2)) == """5""", """value of sum_of_squares(2) is not correct"""

assert str(type(sum_of_squares(10))) == """<class 'int'>""", """type of sum_of_squares(10) is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"""
assert str(sum_of_squares(10)) == """385""", """value of sum_of_squares(10) is not correct"""

assert str(type(sum_of_squares(11))) == """<class 'int'>""", """type of sum_of_squares(11) is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"""
assert str(sum_of_squares(11)) == """506""", """value of sum_of_squares(11) is not correct"""


print('Success!')
```

<!-- grade_id: sum_of_squares_uses_squares, points: 0.5, grade: true, solution: false -->
```python
"""Check that sum_of_squares relies on squares."""

orig_squares = squares
del squares

from hashlib import sha1
assert str(type(test_func_throws(lambda : sum_of_squares(1), NameError))) == """<class 'bool'>""", """type of test_func_throws(lambda : sum_of_squares(1), NameError) is not bool. test_func_throws(lambda : sum_of_squares(1), NameError) should be a bool"""
assert str(test_func_throws(lambda : sum_of_squares(1), NameError)) == """True""", """boolean value of test_func_throws(lambda : sum_of_squares(1), NameError) is not correct"""


squares = orig_squares

print('Success!')
```

---
## Part C (1 point)

Using LaTeX math notation, write out the equation that is implemented by your `sum_of_squares` function.

<!-- grade_id: sum_of_squares_equation, points: 1, grade: true, solution: true -->
$\sum_{i=1}^n i^2$

---
## Part D (2 points)

Find a usecase for your `sum_of_squares` function and implement that usecase in the cell below.

<!-- grade_id: sum_of_squares_application, points: 2, grade: true, solution: true -->
```python
def pyramidal_number(n):
    """Returns the n^th pyramidal number"""
    return sum_of_squares(n)
```

---
## Part E (4 points)

State the formulae for an arithmetic and geometric sum and verify them numerically for an example of your choice.

---
## Part F (1 points)

<!-- grade_id: cell-d3df8cd59fd0eb74, points: 0, grade: false, solution: true -->
```python
my_dictionary = {
    'one'   : 1,
    'two'   : 2,
    'three' : 3
}
```

<!-- grade_id: cell-6e9ff83aa5dfaf17, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(my_dictionary)) == """<class 'dict'>""", """type of my_dictionary is not dict. my_dictionary should be a dict"""
assert str(len(list(my_dictionary.keys()))) == """3""", """number of keys of my_dictionary is not correct"""
assert str(sorted(map(str, my_dictionary.keys()))) == """['one', 'three', 'two']""", """keys of my_dictionary are not correct"""
assert str(sorted(map(str, my_dictionary.values()))) == """['1', '2', '3']""", """correct keys, but values of my_dictionary are not correct"""
assert str(my_dictionary) == """{'one': 1, 'two': 2, 'three': 3}""", """correct keys and values, but incorrect correspondence in keys and values of my_dictionary"""

assert str(type(my_dictionary["one"])) == """<class 'int'>""", """type of my_dictionary[\"one\"] is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"""
assert str(my_dictionary["one"]) == """1""", """value of my_dictionary[\"one\"] is not correct"""

print('Success!')
```