<!-- grade_id: cell-39cee8b588a2b231, points: 0, grade: false, solution: true -->
```python
def twice(x):
    ### BEGIN SOLUTION ###
    return 2*x    
    ### END SOLUTION ###
```

<!-- grade_id: cell-285e58af6409632f, points: 1, grade: true, solution: false -->
```python
assert twice(9)==18
```

Define a function to calculate three times a number!